rule_files:
- ../aggregate-prometheus-rules/fluent-bit.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  external_labels:
    seed: aws
  input_series:
  # FluentBitDown
  - series: 'up{job="fluent-bit"}'
    values: '0+0x30'
  alert_rule_test:
  - eval_time: 15m
    alertname: FluentBitDown
    exp_alerts:
    - exp_labels:
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: "There are no fluent-bit pods running on seed: aws. No logs will be collected."
        summary: Fluent-bit is down

