rule_files:
- kube-etcd3-test.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  input_series:
  # KubeEtcdTestDown
  # KubeEtcdBackupRestoreTestDown
  - series: 'up{job="kube-etcd3-test"}'
    values: '0+0x30 1+0x40'
  # KubeEtcd3TestNoLeader
  - series: 'etcd_server_has_leader{job="kube-etcd3-test"}'
    values: '0+0x30'
  # KubeEtcd3HighNumberOfFailedProposals
  - series: 'etcd_server_proposals_failed_total{job="kube-etcd3-test", pod="etcd"}'
    values: '0+2x60 121+0x60'
  # KubeEtcd3DbSizeLimitApproaching
  # KubeEtcd3DbSizeLimitCrossed
  - series: 'etcd_mvcc_db_total_size_in_bytes{job="kube-etcd3-test"}'
    values: '7194070000+107374182x20' # 6.7GB 6.8GB 6.9GB .. 7.7GB
  # KubeEtcdDeltaBackupFailed
  - series: 'etcdbr_snapshot_latest_timestamp{job="kube-etcd3-backup-restore-test",kind="Incr"}'
    values: '0+0x62'
  - series: 'etcdbr_snapshot_required{job="kube-etcd3-backup-restore-test",kind="Incr"}'
    values: '1+0x62'
  # KubeEtcdFullBackupFailed
  - series: 'etcdbr_snapshot_latest_timestamp{job="kube-etcd3-backup-restore-test",kind="Full"}'
    values: '0+0x2912'
  - series: 'etcdbr_snapshot_required{job="kube-etcd3-backup-restore-test",kind="Full"}'
    values: '1+0x2912'
  # KubeEtcdRestorationFailed
  - series: 'etcdbr_restoration_duration_seconds_count{job="kube-etcd3-backup-restore-test",succeeded="false"}'
    values: '0+0x7 1 2 2'
  # KubeEtcdBackupRestoreTestDown
  - series: 'up{job="kube-etcd3-backup-restore-test"}'
    values: '0+0x60 1+0x10'
  - series: 'etcdbr_snapshotter_failure{job="kube-etcd3-backup-restore-test"}'
    values: '1+1x30 1+0x40'
  alert_rule_test:
  - eval_time: 15m
    alertname: KubeEtcdTestDown
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd3 cluster test is unavailable or cannot be scraped. As long as etcd3 test is down the cluster is unreachable.
        summary: Etcd3 test cluster down.
  - eval_time: 15m
    alertname: KubeEtcd3TestNoLeader
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd3 test has no leader. No communication with etcd test possible. Apiserver is read only.
        summary: Etcd3 test has no leader.
  - eval_time: 1h
    alertname: KubeEtcd3HighNumberOfFailedProposals
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: warning
        type: seed
        visibility: operator
        pod: etcd
        job: kube-etcd3-test
      exp_annotations:
        description: Etcd3 test pod etcd has seen 121 proposal failures within the last hour.
        summary: High number of failed etcd proposals
  - eval_time: 5m
    alertname: KubeEtcd3DbSizeLimitApproaching
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-test
        service: etcd
        severity: warning
        type: seed
        visibility: all
      exp_annotations:
        description: Etcd3 test DB size is approaching its current practical limit of 8GB. Etcd quota might need to be increased.
        summary: Etcd3 test DB size is approaching its current practical limit.
  - eval_time: 10m
    alertname: KubeEtcd3DbSizeLimitCrossed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-test
        service: etcd
        severity: critical
        type: seed
        visibility: all
      exp_annotations:
        description: Etcd3 test DB size has crossed its current practical limit of 8GB. Etcd quota must be increased to allow updates.
        summary: Etcd3 test DB size has crossed its current practical limit.
  - eval_time: 31m
    alertname: KubeEtcdDeltaBackupFailed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-backup-restore-test
        kind: Incr
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: No delta snapshot for the past at least 30 minutes.
        summary: Etcd delta snapshot failure.
  - eval_time: 1456m
    alertname: KubeEtcdFullBackupFailed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-backup-restore-test
        kind: Full
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: No full snapshot taken in the past day.
        summary: Etcd full snapshot failure.
  - eval_time: 5m
    alertname: KubeEtcdRestorationFailed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-backup-restore-test
        succeeded: false
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd data restoration was triggered, but has failed.
        summary: Etcd data restoration failure.
  - eval_time: 16m
    alertname: KubeEtcdBackupRestoreTestDown
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-backup-restore-test
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd backup restore test process down or snapshotter failed with error. Backups will not be triggered unless backup restore is brought back up. This is unsafe behaviour and may cause data loss.
        summary: Etcd backup restore test process down or snapshotter failed with error
  - eval_time: 30m
    alertname: KubeEtcdBackupRestoreTestDown
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd backup restore test process down or snapshotter failed with error. Backups will not be triggered unless backup restore is brought back up. This is unsafe behaviour and may cause data loss.
        summary: Etcd backup restore test process down or snapshotter failed with error
  - eval_time: 35m
    alertname: KubeEtcdBackupRestoreTestDown
    exp_alerts:
