rule_files:
- seed-victoria-logs.prometheusrule.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  external_labels:
    seed: aws
  input_series:
  # VictoriaLogsDown
  - series: 'up{job="victoria-logs"}'
    values: '0+0x30'
  alert_rule_test:
  - eval_time: 30m
    alertname: VictoriaLogsDown
    exp_alerts:
    - exp_labels:
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: "There are no VictoriaLogs pods running on seed: aws. No logs will be collected."
        summary: VictoriaLogs is down
